# Kick-off Assignment

The exercise for this week is meant to help you practice using Python and Jupyter Notebook, and practice saving changes to your files using Git and GitHub. You will also get to practice the Markdown syntax. 

## Notes about the exercise notebook

1. The Python code cells where you should make changes contain text that reads

    ```Python

    # YOUR CODE HERE
    raise NotImplementedError()
    
    ```
    <br/>**You should delete that text and replace it with your solution (code)**.

   
2. The Markdown cells where you should make changes contain text that reads

    ```text
    
    YOUR ANSWER HERE
    ```
    <br/>**You should delete that text and replace it with your answer**.

3. Some of the cells in this Jupyter Notebook are "Read-only", which means that you are not able to edit them.
4. Some of the code cells contain tests. If these tests fail (i.e., you see an error message), you know there is still something wrong with your code that you should fix.

<!-- ### Tips for completing this exercise

- Use **exactly** the same variable names as in the instructions because your answers will be automatically graded, and the tests that grade your answers rely on following the same formatting or variable naming as in the instructions.
- **Please do not**:

    - **Change the file names**. Do all of your editing in the provided `Exercise-2.ipynb` file (this file).
    - **Copy/paste cells in this notebook**. We use an automated grading system that will fail if there are copies of code cells.
    - **Change the existing cell types**. You can add cells, but changing the cell types for existing cells (from code to markdown, for example) will also cause the automated grader to fail. -->

## Problem 1 - Cloning Assignment repository in JupyterLab

Problem 1 involves cloning the kick-off assignment to your local computer by accepting the GitHub Classroom assignment and utilizing the JupyterLab Git extension. To clone the kick-off assignment for editing on your local system, refer to the instructions detailed in the tutorial [Preparations](https://engo645-assignments.readthedocs.io/en/latest/tutorials/L1/prepration.html).


## Problem 2 - Making and changing Numpy Arrays, and using Git

Imagin a hospital tested the age and body fat data for 10 randomly selected adults with the following results. 

| Age      | % Fat     |
| -------- | :-------: |
| 23       | 9.5       |
| 23       | 26.5      |
| 27       | 7.8       |
| 27       | 17.8      |
| 39       | 31.4      |
| 41       | 25.9      |
| 47       | 27.4      |
| 49       | 27.2      |
| 50       | 31.2      |
| 52       | 34.6      |

Your initial task for this assignment involves utilizing the ``` numpy ``` library to create arrays representing the ages and the body fat percentages of these individuals. 
You will then compute the mean, median, and standard deviation for both the age and % fat datasets.


### Part 1 - Creating the arrays

In the cell below you should create two arrays:

- `age_array` should contain the ages values
- `body_fat_array` should contain the %fat values

```python

```

<!-- grade_id: problem_1_ag_p1_station_lists, points: 0, grade: false, solution: true -->
```python
# import the numpy array package   
import numpy as np # type: ignore

# Create the arrays below
age_array = np.array([23,23,27,27,39,41,47,49,50,52])
body_fat_array = np.array([9.5,26.5,7.8,17.8,31.4,25.9,27.4,27.2,31.2,34.6])

```

Let's now check the lists contain the expected values.

<!-- grade_id: problem_1_ag_p1_station_lists_test1, points: 0.5, grade: true, solution: false -->
```python
######### This is a test cell that checks that the arrays lengths are correct #########
# Running this cell should not produce any errors

assert len(age_array) == 10, 'The age_array should have 10 items.'
assert len(body_fat_array) == 10, 'The body_fat_array should have 10 items.'
```

<!-- grade_id: problem_1_ag_p1_station_lists_test2, points: 0.5, grade: true, solution: false -->
```python
######### This is a test cell that checks that the first item in the arrays is correct #########
# Running this cell should not produce any errors

assert age_array[0] == 23, 'The first item in the age_array should be 23'
assert body_fat_array[0] == 9.5, 'The first item in the body_fat_array should be 9.5'
```

### Part 2 - Calculating Mean, Median, and Standard Deviation

Let's now calculate some basic statistics about our dataset.


In the cell below you should use `numpy` build in functions `np.mean()`, `np.median()`, `np.std()`  to calculate mean, median, and standard deviation for each array..

<!-- grade_id: problem_1_ag_p2_modify_lists, points: 0, grade: false, solution: true -->
```python
# Below, add the stations from the second table to the lists

age_mean = np.mean([23,23,27,27,39,41,47,49,50,52])
age_median = np.median([23,23,27,27,39,41,47,49,50,52])
age_std = np.std([23,23,27,27,39,41,47,49,50,52])

body_fat_mean = np.mean([9.5,26.5,7.8,17.8,31.4,25.9,27.4,27.2,31.2,34.6])
body_fat_median = np.median([9.5,26.5,7.8,17.8,31.4,25.9,27.4,27.2,31.2,34.6])
body_fat_std = np.std([9.5,26.5,7.8,17.8,31.4,25.9,27.4,27.2,31.2,34.6])

```

And we can now check to see whether the calculated parameters for "age" are correct

<!-- grade_id: problem_1_ag_p2_modify_lists_test1, points: 0.5, grade: true, solution: false -->
```python
######### This is a test cell that checks that the calculated parameters are correct #########
# Running this cell should not produce any errors

print(f'The mean age is {age_mean:.2f}')
print(f'The median age is {age_median:.2f}')
print(f'The standard deviation of age is {age_std:.2f}')
```

Output:
```
The mean age is 37.80
The median age is 40.00
The standard deviation of age is 11.15
```

And we can now check to see whether the calculated parameters for "%fat" are correct

<!-- grade_id: problem_1_ag_p2_modify_lists_test2, points: 0.5, grade: true, solution: false -->
```python
######### This is a test cell that checks that the calculated parameters are correct #########
# Running this cell should not produce any errors

print(f'The mean body fat is {body_fat_mean:.2f}')
print(f'The median body fat is {body_fat_median:.2f}')
print(f'The standard deviation of body fat is {body_fat_std:.2f}')
```

Output:
```
The mean body fat is 23.93
The median body fat is 26.85
The standard deviation of body fat is 8.73
```

### Part 3 - Python + Google (optional)

This section is optional. In Python, there's a powerful library called `Pandas` that facilitates working with data in the form of `DataFrame`. Dataframes are particularly useful for handling tables, such as Excel files or CSV data. If you'd like, you're free to do some Googling to learn how to create our age and %fat table using a Pandas DataFrame. Additionally, try calculating the mean, median, and standard deviation of the data using Pandas' functionalities.

**NOTE**: If you plan to skip this optional problem, be sure to remove the `raise NotImplementedError()` line in the cell below so that you don't encounter any errors when running your notebook!


<!-- grade_id: problem_1_mg_p5_zip_test, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
import pandas as pd

# Create the data
data = {
    "Age": [23, 23, 27, 27, 39, 41, 47, 49, 50, 52],
    "% Fat": [9.5, 26.5, 7.8, 17.8, 31.4, 25.9, 27.4, 27.2, 31.2, 34.6]
}

# Create a pandas DataFrame
df = pd.DataFrame(data)

# Display the DataFrame
print(df)

# raise NotImplementedError()
```

Output:
```
Age  % Fat
0   23    9.5
1   23   26.5
2   27    7.8
3   27   17.8
4   39   31.4
5   41   25.9
6   47   27.4
7   49   27.2
8   50   31.2
9   52   34.6
```

## Problem 3 - Practice using Markdown (2 points)

The last task in this assigmnet is to answer some questions and add an image in this notebook using Markdown. **NOTE**: You may want to read a bit more about [formatting text in Github-flavored Markdown](https://help.github.com/articles/basic-writing-and-formatting-syntax/).



### Part 1 - Questions about this assigments

Please use Markdown to give your responses to the three questions below about this assigmnets.

- What did you learn?
- What was unclear?
- What would you change about the tutorials?
    
*Please use a Markdown list when answering these questions.*

<!-- grade_id: problem_3_mg_p1_questions, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

# What I learnt?
I learnt how to clone a repository.
I learnt how to set up python in Anaconda Navigator.
I learnt how to write python code using numpy and Pandas.

# What was Unclear?
How to use VScode for python coding.
How to Clone a github repository in google collab environment.
How to push and pull from VScode to Repository.
I installed Anaconda but Jupyter lab did not lauch despite several reboot. How do I fix this?

# What I would change about the tutorial?

Adding the section on how to clone my github repository using google collab.



### Part 2 - Adding an image

Let's practice adding an image using markdown! You can search for an interesting picture related to spatial data mining on Google. Write a short description explaining what the image shows and its importance in understanding spatial data mining.

You can add the image by using a URL or, for an extra challenge, upload an image to your GitHub repository and link it to this notebook.


![Sample Image](https://www.dbs.ifi.lmu.de/Forschung/KDD/SpatialKDD/)

Data Mining from Spatial Data